<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Create a object that represents a store. Have the initialization method take the square footage of the store. Create a method that can calculate the electrical costs for the store given:

- $0.15 per kWh
- kWh per square foot per hour of 0.19

Ensure your class has the following API:

Instance attributes: 
- `square_footage`

Class level attributes:
- DOLLARS_PER_KWH
- KWH_PER_SQUARE_FOOTAGE_PER_HOUR

The method that calculates the electrical cost for the store should be named `electrical_cost`.

[2 points]

<!-- grade_id: cell-e0ff639912440f53, points: 0, grade: false, solution: true -->
```python
class Store:
    
    DOLLARS_PER_KWH = 0.15
    KWH_PER_SQUARE_FOOTAGE_PER_HOUR = 0.19
    
    def __init__(self,square_footage):
        self.square_footage = square_footage
        
    def electrical_cost(self,hour):
        self.electrical_cost=self.square_footage * hour * Store.KWH_PER_SQUARE_FOOTAGE_PER_HOUR * Store.DOLLARS_PER_KWH
        return self.electrical_cost
        
```

<!-- grade_id: cell-78899230c51c73e2, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert not hasattr(Store, 'square_footage')
### END TESTS
```

<!-- grade_id: cell-e4863bba06d08e3f, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store(2000).square_footage == 2000
### END TESTS
```

<!-- grade_id: cell-47704c4e3e538b07, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert hasattr(Store, 'DOLLARS_PER_KWH')
### END TESTS
```

<!-- grade_id: cell-72b0af5f335f9ade, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert hasattr(Store, 'KWH_PER_SQUARE_FOOTAGE_PER_HOUR')
### END TESTS
```

<!-- grade_id: cell-da6780d986482482, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store.DOLLARS_PER_KWH == 0.15
### END TESTS
```

<!-- grade_id: cell-094c43f4a0ba194f, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store.KWH_PER_SQUARE_FOOTAGE_PER_HOUR == 0.19
### END TESTS
```

<!-- grade_id: cell-32758bf2fddd6948, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert callable(Store(2000).electrical_cost)
### END TESTS
```

<!-- grade_id: cell-e9ab2e3e055a711f, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store(2000).electrical_cost(900) == 51300
### END TESTS
```

**Q**: Create a counter class that takes a string and can produce (via a method named `count`) a `dict` containing each word in the string as its keys and the counts of each word in the string as its values. Words are any token separated by a space. The class should not be case sensitive (ie. "It" and "it" are the same word).


```python
>>> c = Counter('This is my best string. It is a great string.')
>>> c.count()
{'a': 1,
 'best': 1,
 'great': 1,
 'is': 2,
 'it': 1,
 'my': 1,
 'string.': 2,
 'this': 1}
 ```
 
**Class API**

Instance attributes:
  - "string"
  
You can not use the collections module.
 
[2 points]

<!-- grade_id: cell-0be5d346f805f3fb, points: 0, grade: false, solution: true -->
```python
class Counter:
    
    def __init__(self,string):
        self.string=string
        
    def count(self):
        dic = dict()
        lower_string = self.string.lower()
        words = lower_string.split()
        
        for word in words:
            if word in dic:
                dic[word] = dic[word] + 1
            else:
                dic[word] = 1
                
        self.dic = dic
        
        return self.dic
```

<!-- grade_id: cell-59639940597f3439, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert len(counter.count()) == 8
### END TESTS
```

<!-- grade_id: cell-a8f67550bdfc6edd, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('')
assert len(counter.count()) == 0
### END TESTS
```

<!-- grade_id: cell-ea0b8e76389a0ae2, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert hasattr(counter, 'string')
### END TESTS
```

<!-- grade_id: cell-9aa8b06a3d7baaaf, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert not hasattr(Counter, 'string')
### END TESTS
```

<!-- grade_id: cell-91fa2691c9bea2eb, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
counter = Counter('What what what? What is this? Do you know what this is?')
result = counter.count()
expected = {
    'do': 1,
    'is': 1,
    'is?': 1,
    'know': 1,
    'this': 1,
    'this?': 1,
    'what': 4,
    'what?': 1,
    'you': 1
}
assert result == expected
### END HIDDEN TESTS
```

<!-- grade_id: cell-b3934871fd7f7a25, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
counter = Counter('What what what? What is this? Do you know what this is?')
result = counter.count()
assert result['what'] == 4
assert result['what?'] == 1
### END HIDDEN TESTS
```

<!-- grade_id: cell-fbb2ec3c367fafd9, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
counter = Counter('What what what? What is this? Do you know what this is?')
result = counter.count()
assert 'Do' not in result
### END HIDDEN TESTS
```

<!-- grade_id: cell-69816282ea2b630d, points: 0.25, grade: true, solution: false -->
```python
 ### BEGIN HIDDEN TESTS
counter = Counter('What what what? What is this? Do you know what this is?')
result = counter.count()
assert ' ' not in result
### END HIDDEN TESTS
```

**Q**: Update the counter class to have the appropriate methods for presentation when being passed to `repr` and `str`. `repr` should return the object class name and a truncated version of the counter's string. The string should be truncated to 10 characters and an ellipsis added if the string is over 10 characters. `str` should return the object class name and the key count.

You may need to search the Python documentation for the correct methods to implement for this problem.

```python
>>> obj = Counter('Fe Fi Fo Fum, where is my drum')
>>> repr(obj)
<Counter 'Fe Fi Fo F...'>
>>> str(obj)
'Counter(8)'
```

[1 point]

<!-- grade_id: cell-eff92fe808b528d9, points: 0, grade: false, solution: true -->
```python
class Counter:
    
    def __init__(self,string):
        self.string=string
        
    def count(self):
        dic = dict()
        lower_string = self.string.lower()
        words = lower_string.split()
        for word in words:
            if word in dic:
                dic[word] = dic[word] + 1
            else:
                dic[word] = 1
        self.dic = dic
        return self.dic
    
    def __repr__(self):
        if len(self.string)>10:
            return f"<Counter '{self.string[0:10]+'...'}'>"
        else:
            return f"<Counter '{self.string}'>"
    
    def __str__(self):
        return f'Counter({len(Counter.count(self))})'
```

<!-- grade_id: cell-42a5d9c016cae56f, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert hasattr(Counter, '__repr__')
assert hasattr(Counter, '__str__')
### END HIDDEN TESTS
```

<!-- grade_id: cell-76d7e759c0c50b24, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert repr(counter) == "<Counter 'This is my...'>"
### END TESTS
```

<!-- grade_id: cell-a6990dafdb2b143c, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('Woot woot!')
assert repr(counter) == "<Counter 'Woot woot!'>"
### END TESTS
```

<!-- grade_id: cell-d5d109b50b154b18, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert str(counter) == 'Counter(8)'
### END TESTS
```

**Q**: Create a custom exception for a severe lack of parties. It should be called `SevereLackOfPartiesError` and should subclass `ValueError`. The class should have a docstring that describes what the error represents. Use your imagination for the exact text. 

[1 point]

<!-- grade_id: cell-7368367ae633aef9, points: 0, grade: false, solution: true -->
```python
class SevereLackOfPartiesError(ValueError):
    
    ''' Used to indicate attributes that does not make any sense which may be caused by the writer's severe lack of parties. '''
```

<!-- grade_id: cell-78f404be90d25d2e, points: 0.5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert issubclass(SevereLackOfPartiesError, ValueError)
### END HIDDEN TESTS
```

<!-- grade_id: cell-1b63d2e6166363a1, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert SevereLackOfPartiesError.__doc__
### END HIDDEN TESTS
```

<!-- grade_id: cell-ebc6c2afbc97b842, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert len(SevereLackOfPartiesError.__doc__) > 0
### END HIDDEN TESTS
```

**Q**: Create a class `Str` that subclasses `str`. Add a method to the subclass that checks if the string does not start with a given string. Call that method `notstartswith`. This class should only take up 3 lines of code however you will not be graded on the length of your class.

[1 point]

<!-- grade_id: cell-931411951b8973e6, points: 0, grade: false, solution: true -->
```python
class Str(str):
    
    def __init__(self, content):
        self.content=str(content)
        
    def notstartswith(self,given_string):
        return not self.content.startswith(given_string)
```

<!-- grade_id: cell-dde9e6cf99caf6bc, points: 0.34, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert issubclass(Str, str)
### END HIDDEN TESTS
```

<!-- grade_id: cell-767d16e4e70c6601, points: 0.33, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert Str('This string does not ...').notstartswith('Hi')
### END HIDDEN TESTS
```

<!-- grade_id: cell-7f7773df3332732e, points: 0.33, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert not Str('This string does not ...').notstartswith('This')
### END HIDDEN TESTS
```

**Q**: Create a `BankAccount` class that is instantiated with an account number and an optional initial balance. Add methods to withdraw and deposit to the account's balance. Also, add a method that allows for the casting of objects of this class to a string.

```python
a = BankAccount(12345678, balance=10)
a.deposit(5)
a.withdraw(3)
str(a)

'Account ID: 12345678 - Balance: 12'
```

[2 points]

<!-- grade_id: cell-31e546afd67a9412, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    
    def __init__(self,accountNumber, **kwargs):
        self.accountNumber = accountNumber
        if kwargs == dict():
            self.balance=0
        else:
            self.balance = kwargs.get('balance')
        
    def deposit(self, money):
        self.balance += money
        
    def withdraw(self, money):
        self.balance -= money
        
    def __str__(self):
        return f'Account ID: {self.accountNumber} - Balance: {self.balance}'
```

<!-- grade_id: cell-78f78a2ea5d08b99, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert str(BankAccount(4839283, balance=2010)) == 'Account ID: 4839283 - Balance: 2010'
### END TESTS
```

<!-- grade_id: cell-4de2527d92d58164, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert hasattr(BankAccount, '__str__')
### END TESTS
```

<!-- grade_id: cell-e706bb44eae7db75, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert BankAccount(1).balance == 0
### END TESTS
```

<!-- grade_id: cell-64c5f7b45a14a431, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert hasattr(BankAccount(1), 'balance')
### END TESTS
```

<!-- grade_id: cell-d69d7f2056853513, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert not hasattr(BankAccount, 'balance')
### END TESTS
```

<!-- grade_id: cell-da11f9aeadad2c39, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
a = BankAccount(1)
a.deposit(150)
assert a.balance == 150
### END TESTS
```

<!-- grade_id: cell-0f6bc2f6bd8c7a9c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = BankAccount(1, balance=150)
a.withdraw(27)
assert a.balance == 123
### END TESTS
```

**Q**: Subclass the `BankAccount` class from the previous question to create a `MinBalanceBankAccount` class. This class should require a minimum balance at all times and should raise an exception (of the type `ValueError`) when an attempt to drop the balance below the minumum occurs. The `MinBalanceBankAccount` should be instantiated with the account number, inital balance, and the minimum balance. No error checking against the minimum balance is necessary upon instantiation.

```python
a = MinBalanceBankAccount(account_id=1, balance=100, minimum_balance=10)
a.withdraw(100)
...
ValueError: Sorry, minimum balance must be maintained.
```

If raised, the `ValueError` should have the message: 'Sorry, minimum balance must be maintained.'

*A minimum of extra code should be implemented* in `__init__` and `withdraw` to meet the requirements of this question. Ie. use `super`!

[2 points]

<!-- grade_id: cell-ffede70c71976fef, points: 0, grade: false, solution: true -->
```python
class MinBalanceBankAccount(BankAccount):
    
    def __init__(self, accountNumber, inital_balance, min_balance):
        super().__init__(accountNumber)
        self.min_balance = min_balance
        self.balance = inital_balance
        
    def withdraw(self, money):
        if self.balance - money < self.min_balance:
            raise ValueError('Sorry, minimum balance must be maintained.')
        self.balance -=  money
```

<!-- grade_id: cell-21fda7b0a6a440d5, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
try:
    MinBalanceBankAccount(1, 100, 10).withdraw(100)
    assert False, 'Should have raised value Error'
except ValueError as e:
    pass
### END TESTS
```

<!-- grade_id: cell-4d189116d6cb3d51, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
try:
    MinBalanceBankAccount(1, 100, 10).withdraw(100)
    assert False, 'Should have raised value Error'
except ValueError as e:
    assert str(e) == 'Sorry, minimum balance must be maintained.'
### END TESTS
```

<!-- grade_id: cell-496de9ecb2b38244, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert 'super' in MinBalanceBankAccount.__dict__['__init__'].__code__.co_names
### END HIDDEN TESTS
```

<!-- grade_id: cell-2d6c5e93ce362aa9, points: 0.25, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert 'super' in MinBalanceBankAccount.__dict__['withdraw'].__code__.co_names
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-1-2373a11adcc7> in <module>
      1 ### BEGIN HIDDEN TESTS
----> 2 assert 'super' in MinBalanceBankAccount.__dict__['withdraw'].__code__.co_names
      3 ### END HIDDEN TESTS

AssertionError:
```

<!-- grade_id: cell-74b6dba911c7df84, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert issubclass(MinBalanceBankAccount, BankAccount)
### END TESTS
```

<!-- grade_id: cell-3d7deeb06f900ce3, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = MinBalanceBankAccount(1, 100, 10)
a.deposit(150)
a.withdraw(27)
assert a.balance == 223
### END TESTS
```

**Extra Practice** Create a class that produces objects that are both iterables and iterators. 

<!-- grade_id: cell-acfb6cbf62089d2f, points: 0, grade: false, solution: true -->
```python
class MyRange:
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
File "<ipython-input-1-559c84af744e>", line 3
    raise NotImplementedError()
    ^
IndentationError: expected an indented block
```

<!-- grade_id: cell-ef5f590e0182d168, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
assert list(MyRange(1, 5)) == [1, 2, 3, 4], list(MyRange(1, 5))
### END TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-68c70b891a16> in <module>
      1 ### BEGIN TESTS
----> 2 assert list(MyRange(1, 5)) == [1, 2, 3, 4], list(MyRange(1, 5))
      3 ### END TESTS

NameError: name 'MyRange' is not defined
```

```python

```